```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

<table style="width: 100%;">
<tr style="background-color: transparent;">
<td width="100px"><img src="https://cs104williams.github.io/assets/cs104-logo.png" width="90px" style="text-align: center"/></td>
<td>
  <p style="margin-bottom: 0px; text-align: left; font-size: 18pt;"><strong>CSCI 104: Data Science and Computing for All</strong><br>
                Williams College<br>
                Fall 2026</p>
</td>
</tr>


# Lab 3: Table Manipulation and Visualization

<hr style="margin: 0px; border: 3px solid #500082;"/>

<h2>Instructions</h2>

- Before you begin, execute the cell at the TOP of the notebook to load the provided tests, as well as the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute these cells again.  
- Be sure to consult your [Python Reference](https://cs104williams.github.io/assets/python-library-ref.html)!
- Complete this notebook by filling in the cells provided. For problems asking you to write explanations, you **must** provide your answer in the designated space. 
- Please be sure to not re-assign variables throughout the notebook.  For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously.
- This lab has hidden tests on it. That means even though tests may say 100% passed, doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the lab.
- To use one or more late days on this lab, please fill out our [late day form](https://forms.gle/4sD16h3hN1xRqQM27) **before** the due date.

<hr/>
<h2>Setup</h2>


```python
# Run this cell to set up the notebook.
# These lines import the numpy, datascience, and cs104 libraries.

import numpy as np
from datascience import *
from cs104 import *
%matplotlib inline
```

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 1. Unemployment Data (40 pts)



<font color='#B1008E'>
    
##### Learning objectives
- Sort data in tables
- Select and manipulate rows and columns in tables in order to answer quantitative questions 
- Create and analyze line plots
</font>

The Federal Reserve Bank of St. Louis publishes data about jobs in the US.  Below, we've loaded data on unemployment in the United States, one row per quarter (the last month of each quarter), from 1994 through the middle of 2026. There are many ways of defining unemployment, and our dataset includes two notions of the unemployment rate:

1. **NEI:** Among people who are able to work and are looking for a full-time job, the percentage who can't find a job.  This is called the Non-Employment Index, or NEI.
2. **NEI-PTER:** Among people who are able to work and are looking for a full-time job, the percentage who can't find a job *or* are only working at a part-time job.  The latter group is called "Part-Time for Economic Reasons", or PTER.  The acronym for this index of unemployment is NEI-PTER. 

The source of the data is [here](https://fred.stlouisfed.org/categories/33509) (series `NEIM156SFRBRIC` and `NEIPTERM156SFRBRIC`, computed by the Federal Reserve Bank of Richmond).

#### Part 1.1 (5 pts)


 The data are in a CSV file called `unemployment.csv`.  Load that file into a table called `unemployment`. 



```python
unemployment = ...
unemployment
```

```python
grader.check("p1.1")
```

#### Part 1.2 (5 pts)


 Sort the data in descending order by NEI, naming the sorted table `by_nei`.  Create another table called `by_nei_pter` that's sorted in descending order by NEI-PTER instead. 



```python
by_nei = ...
by_nei_pter = ...
```

```python
grader.check("p1.2")
```

```python
# Run this cell to check your by_nei table. You do not need to change the code.
by_nei.show(5)
```

```python
# Run this cell to check your by_nei_pter table. You do not need to change the code.
by_nei_pter.show(5)
```

#### Part 1.3 (5 pts)


 Use `take` to make a table containing the data for the 10 quarters when NEI was greatest.  Call that table `greatest_nei`.

`greatest_nei` should be sorted in descending order of `NEI`. Note that each row of `unemployment` represents a quarter. 

```python
greatest_nei = ...
greatest_nei
```

```python
grader.check("p1.3")
```

#### Part 1.4 (5 pts)


Economists care about PTER (recall: "Part-Time for Economic Reasons") for many reasons. For example, it's believed that many people became PTER in the "Great Recession" of 2008-2009, and again during the COVID-19 pandemic in 2020.   
 
Compute an array containing the percentage of people who were PTER in each quarter.  (The first element of the array should correspond to the first row of `unemployment`, and so on.) 

*Notes:* 
- Use the original `unemployment` table for this.
- Recall, NEI is the percentage of people who are unemployed; and NEI-PTER is the percentage of people who are unemployed (included in the NEI) plus the percentage of people who are PTER. 

As a sanity check, the quarterly PTER you compute should be less than 2% for all of the quarters in this data.

```python
pter = ...
pter
```

```python
grader.check("p1.4")
```

#### Part 1.5 (5 pts)


Extend the `unemployment` table with a new column named `PTER` that contains the values in the `pter` array.  Sort the resulting table by that column in descending order.  Call the resulting table `by_pter`.

```python
by_pter = ...
by_pter
```

```python
grader.check("p1.5")
```

<!-- BEGIN QUESTION -->

#### Part 1.6 (5 pts)


Create a line plot of PTER over time. 

To do this, create a new table called `pter_over_time` that adds the `year` array and the `pter` array to the `unemployment` table. Label these columns `Year` and `PTER`. Then, generate a line plot using one of the table methods you've learned in class.

The order of the columns matter for our correctness tests, so be sure `Year` comes before `PTER`. 

*Clarification: When constructing `pter_over_time`, do not just add the `year` column to the `by_pter` table. Please follow the directions in the question above.*

```python
year = 1994 + np.arange(by_pter.num_rows)/4 #this accounts for the quarters 
pter_over_time = ...
plot = ...

# This line will sets the bounds on the y-axis
plot.set_ylim(0,2) 
```

```python
grader.check("p1.6")
```

<!-- END QUESTION -->

#### Part 1.7 (5 pts)


 Were PTER rates high during the Great Recession (that is to say, were PTER rates particularly high in the years 2008 through 2011)? Examine the line plot you just created. Then assign `highPTER` to `True` if you think PTER rates were high in this period, or `False` if you think they weren't. 

```python
highPTER = ...
```

```python
grader.check("p1.7")
```

#### Part 1.8 (5 pts)


The plot shows two bumps in PTER: a long, broad one during the Great Recession, and a short, sharp one in 2020, when the COVID-19 pandemic shut down much of the economy.  Let's compare how long each one lasted.

Using `pter_over_time`, count the number of quarters in which PTER was above 1.5% during the years 2008 through 2011, and assign that number to `great_recession_quarters`.  Then do the same for the years 2020 through 2021 and assign the result to `covid_quarters`.

*Hint:* The `Year` column of `pter_over_time` contains numbers like 2008.0, 2008.25, 2008.5, and 2008.75 for the four quarters of 2008, so `are.between(2008, 2012)` selects every quarter from 2008 through 2011.

```python
high_pter = pter_over_time.where("PTER", are.above(1.5))
great_recession_quarters = ...
covid_quarters = ...
print("Quarters with PTER above 1.5%: Great Recession =", great_recession_quarters, ", COVID-19 =", covid_quarters)
```

```python
grader.check("p1.8")
```

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 2. Birth Rates (35 pts)



<font color='#B1008E'>
    
##### Learning objectives
- Use built-in Table functions, Python functions, and numpy functions to answer quantitative questions about a dataset
- Create visualizations to analyze trends in data
</font>

The following table gives census-based population estimates for each state on both July 1, 2023 and July 1, 2024. The last four columns describe the components of the estimated change in population during this time interval. **For all questions below, assume that the word "states" refers to all 52 rows including Puerto Rico & the District of Columbia.**
- The data was taken from [here](https://www2.census.gov/programs-surveys/popest/datasets/2020-2024/state/totals/NST-EST2024-ALLDATA.csv) and cleaned up. 
- If you want to read more about the different column descriptions, click [here](https://www2.census.gov/programs-surveys/popest/technical-documentation/file-layouts/2020-2024/NST-EST2024-ALLDATA.pdf).

```python
pop = Table().read_table("nst-est2024-clean.csv")
pop.show(5)
```

#### Part 2.1 (5 pts)


Assign the variable `us_birth_rate` to the total US annual birth rate (across all states) during this time interval. The annual birth rate for a year-long period is the total number of births in that period as a proportion of the population size at the start of the time period. 

*Hint:* Which year corresponds to the start of the time period?



```python
us_birth_rate = ...
us_birth_rate
```

```python
grader.check("p2.1")
```

#### Part 2.2 (5 pts)


Create a table `movers` which contains *at minimum* the same information as `pop` but only includes states for which the **absolute value** of the **annual rate of migration** was higher than 1%. 
 
- The `MIGRATION` column contains estimated annual net migration counts by state. 
- The annual rate of migration for a year-long period is the net number of migrations (in and out) as a proportion of the population size at the start of the period. You may wish to create a temporary table `migration_rates` that contains the migration rates for each state.

```python
# First, create a new table with the data in pop, plus migration rates.
migration_rates = ...
movers = ...
movers
```

```python
grader.check("p2.2")
```

#### Part 2.3 (5 pts)


Assign `west_births` to the total number of births that occurred in region 4 (the Western US). 

```python
west_births = ...
west_births
```

```python
grader.check("p2.3")
```

#### Part 2.4 (5 pts)


 

In the next question, you will be creating a visualization to understand the relationship between birth and death rates. The annual death rate for a year-long period is the total number of deaths in that period as a proportion of the population size at the start of the time period.

*Note:* We will comparing *each* of the 52 states' birth rate and death rate (*not* the overall birth and death rate for the entire population). 

What visualization is most appropriate to see if there is an association between birth and death rates during a given time interval? 

1. Line Graph
<br>
2. Bar Chart
<br>
3. Scatter Plot

Assign `visualization` below to the number corresponding to the correct visualization. 



```python
visualization = ...
```

```python
grader.check("p2.4")
```

<!-- BEGIN QUESTION -->

#### Part 2.5 (5 pts)


 In the code cell below, create a visualization based on your answer above in order to determine if there is an association between birth rate and death rate during this time interval. It may be helpful to create an intermediate table here. 

```python
# In this cell, use birth_rates and death_rates to generate your chart
birth_rates = ...
death_rates = ...
...
```

<!-- END QUESTION -->

#### Part 2.6 (5 pts)


True or False: There is an association between birth rate and death rate during this time interval. 

Assign `assoc` to `True` or `False` in the cell below.  Base your conclusion on your inspection of the visualization you created for the last part.  Reasonable people can disagree about this one, so we will not grade your answer -- but you will explain it in the next part.

```python
assoc = ...
```

```python
grader.check("p2.6")
```

<!-- BEGIN QUESTION -->

#### Part 2.7 (5 pts)


In a sentence or two, explain your answer to the previous part.  What does the scatter plot show about the relationship between birth rate and death rate across states, and how strong is that relationship?  Are there any states that stand out?

<hr style="margin:0; border: 1px solid #FFBE0A;"/><font color='#FFBE0A'>Written Answer:</font>

_Type your answer here, replacing this text._


<hr style="margin:0; border: 1px solid #FFBE0A;"/>

<!-- END QUESTION -->

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 3. Bluebikes data (30 pts)



<font color='#B1008E'>
    
##### Learning objectives
- Use built-in Table functions, Python functions, and numpy functions to answer quantitative questions about a dataset
- Create histograms from data and understand the histograms' quantitative significance
</font>

Below we load a table containing a random sample of 20,000 rides taken in June 2026 on [Bluebikes](https://bluebikes.com/), the bike-share system in the Boston area.  Riders come in two kinds: **members**, who pay for a monthly or annual pass, and **casual** riders, who buy a single ride or a day pass.
- The `Start Time`, `Start Station`, and `End Station` columns describe where and when each ride began and ended.
- The `Rider Type` column is either `member` or `casual`.
- The `Duration (min)` column contains the length of the ride, **in minutes**.

We split the table into two smaller tables, one for each kind of rider.

```python
bluebikes = Table.read_table("bluebikes.csv")
member = bluebikes.where("Rider Type", are.equal_to("member"))
casual = bluebikes.where("Rider Type", are.equal_to("casual"))
print("Member Table")
member.show(4)
print("Casual Table")
casual.show(4)
```

<!-- BEGIN QUESTION -->

#### Part 3.1 (5 pts)


Produce a histogram that visualizes the distribution of all ride durations by members using the given bins.  Set the plot's title to "Member Rides" by storing your plot in the `member_plot` variable and then using the [`set_title` method](https://cs104williams.github.io/assets/python-library-ref.html#set-labels) for `member_plot`.

```python
equal_bins = np.arange(0, 120, 5)
member_plot = ...
...
```

```python
grader.check("p3.1")
```

<!-- END QUESTION -->

#### Part 3.2 (5 pts)


According to your histogram, what percent of member rides are between 20 and 25 minutes?  *Note:* Recall that the x-axis units are minutes

```python
rides_between_20_and_25 = ...
```

```python
grader.check("p3.2")
```

<!-- BEGIN QUESTION -->

#### Part 3.3 (5 pts)


Now, produce a histogram that visualizes the distribution of all ride durations by casual riders using the given bins.

Set the y-axis of your casual-rider histogram to be between 0 and 7%. This will make it easier to compare your member and casual plots. To do that, store the result of making your histogram in a local variable `casual_plot` and then use the `casual_plot`'s [`set_ylim` method](https://cs104williams.github.io/assets/python-library-ref.html#set-limits).  

Set your plot's title to be "Casual Rides" with the `set_title` method for plots.

```python
equal_bins = np.arange(0, 120, 5)
casual_plot = ...
...
...
```

```python
grader.check("p3.3")
```

<!-- END QUESTION -->

#### Part 3.4 (5 pts)


Let's *manually* compare and inspect these member and casual histograms. 

Assign `member_under_15` and `casual_under_15` to the percentage of rides that are less than 15 minutes for each kind of rider. 
 
 To help you we have provide the "height" of various bars from the histogram you created above. Use the height variables that we have defined for you below in order to compute the percentages. Your solution should only use height variables, numbers, and mathematical operations. You should **not** access the tables `member` and `casual` in any way.  
 
 *Hint:* How does the *width* of the bars in the histogram factor into your answer?

```python
member_under_5_height = 3.8
casual_under_5_height = 1.4
member_5_to_under_10_height = 6.5
casual_5_to_under_10_height = 4.5
member_10_to_under_15_height = 4.1
casual_10_to_under_15_height = 4.1

member_under_15 = ...
casual_under_15 = ...

member_under_15, casual_under_15
```

```python
grader.check("p3.4")
```

#### Part 3.5 (5 pts)


Manually examine the histogram for the distribution of ride durations by members. You can use this histogram to find the median ride duration for members. 

Assign `member_median_bin` to an integer (1, 2, 3, 4, or 5) that corresponds to the bin that contains the median time. 

1. 0-5 minutes  
2. 5-10 minutes  
3. 10-15 minutes  
4. 15-20 minutes  
5. 20-25 minutes  

*Hint:* The median of a sorted list has half of the list elements to its left, and half to its right.  Your previous answer may help you determine where the median is.

```python
member_median_bin = ...
member_median_bin
```

```python
grader.check("p3.5")
```

<!-- BEGIN QUESTION -->

#### Part 3.6 (5 pts)


Identify one difference between the histograms, in terms of actual histogram properties. Then explain why the distributions for members and casual riders are different: who do you think members and casual riders are, and how does that explain the shapes of the two histograms?

*Hint:* Think about why someone would buy an annual pass versus a single ride, and what kinds of trips each group is likely to take.



<hr style="margin:0; border: 1px solid #FFBE0A;"/><font color='#FFBE0A'>Written Answer:</font>

_Type your answer here, replacing this text._


<hr style="margin:0; border: 1px solid #FFBE0A;"/>

<!-- END QUESTION -->

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 4. Iris Varieties (10 pts)



<font color='#B1008E'>
    
##### Learning objectives
- Create plots to aid in exploring a new data set
- Complete a more open-ended data visualization and analysis task
</font>

<!-- BEGIN QUESTION -->

#### Part 4.1 (10 pts)


The Iris dataset is a famous data set dating back to 1936.  It contains the sepal and petal lengths and widths for fifty flowers of three different varieties of irises (a type of flower).  Sepals and petals are identified in the figure below for the three varieties.  We provide an extra file `iris.csv` containing this data set in the lab this week.  All measurements are in cm.


<img src="iris.png" width=100%>

This question is *very* open ended. Your task is simply this: perform exploratory data analysis on this data set.

Giving such an open-ended problem gives you maximum flexibility for thinking and exploration, and also to simulate "real life." 
To give you an idea of our expectations, you should try (and display) several types of plots to explore the data. Consider, for example, whether there are corrleations between sepal/petal length and width, and between sepal and petal dimensions.  Are some varieties more uniform in their petal/sepal dimensions?  Are some varieties generally larger than others?  You should comment on your approach and what you learn from the plots.  **Also, be sure each plot has appropriate titles and labels.**

You may do as little or as much as you like in answering this question, but drawing perhaps around three interesting conclusions about the iris varieties will be sufficient to earn full credit.  More importantly, have fun exploring the data!

*Note:* We give you one "Markdown" cell for text and one code cell for your Python code below.  You may add as many additional cells as you like.  Just click the `+` sign in this tab's toolbar to insert a new cell.  Then select "Markdown" from the toolbar's popup menu if you want the new cell to be for text, or "Code" if you want it to be for Python code.  Markdown cells can include basic formatting.  Click on any of our formatted text cells to see how to create lists, bold text, etc., or have a look [here](https://www.markdownguide.org/basic-syntax).

<hr style="margin:0; border: 1px solid #FFBE0A;"/><font color='#FFBE0A'>Written Answer:</font>

_Type your answer here, replacing this text._


<hr style="margin:0; border: 1px solid #FFBE0A;"/>

```python
# This is a code cell where you can get started on your work.
```

<!-- END QUESTION -->

<hr class="m-0" style="border: 3px solid #500082;"/>

# You're Done!
Follow these steps to submit your work:
* Run the tests and verify that they pass as you expect. 
* Choose **Save Notebook** from the **File** menu.
* **Run the final cell** and click the link below to download the zip file. 

Once you have downloaded that file, go to [Gradescope](https://www.gradescope.com/) and submit the zip file to 
the corresponding assignment. For Lab N, the assignment will be called "Lab N Autograder".

Once you have submitted, your Gradescope assignment should show you passing all the tests you passed in your assignment notebook.


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for p1.1 (points: 5)
```
>>> check_str('unemployment.select(\'Date\', \'NEI\', \'NEI-PTER\').column(\'Date\').item(0) == \'1994-03-01\'', locals())
```
failure_message: check(unemployment.select('Date', 'NEI', 'NEI-PTER').column('Date').item(0) == '1994-03-01')
success_message: check(unemployment.select('Date', 'NEI', 'NEI-PTER').column('Date').item(0) == '1994-03-01')
```
>>> check_str('unemployment.num_rows == 130', locals())
```
failure_message: check(unemployment.num_rows == 130)
success_message: check(unemployment.num_rows == 130)


Tests for p1.2 (points: 5)
```
>>> check_str('by_nei.column(\'Date\').item(0) == \'2020-06-01\'', locals())
```
failure_message: check(by_nei.column('Date').item(0) == '2020-06-01')
success_message: check(by_nei.column('Date').item(0) == '2020-06-01')
```
>>> check_str('by_nei_pter.column(\'Date\').item(0) == \'2020-06-01\'', locals())
```
failure_message: check(by_nei_pter.column('Date').item(0) == '2020-06-01')
success_message: check(by_nei_pter.column('Date').item(0) == '2020-06-01')


Tests for p1.3 (points: 5)
```
>>> check_str('greatest_nei.num_rows == 10', locals())
```
failure_message: check(greatest_nei.num_rows == 10)
success_message: check(greatest_nei.num_rows == 10)
```
>>> check_str('greatest_nei.column(\'Date\').item(0) == \'2020-06-01\'', locals())
```
failure_message: check(greatest_nei.column('Date').item(0) == '2020-06-01')
success_message: check(greatest_nei.column('Date').item(0) == '2020-06-01')


Tests for p1.4 (points: 5)
```
>>> check_str('len(pter) == 130', locals())
```
failure_message: check(len(pter) == 130)
success_message: check(len(pter) == 130)
```
>>> check_str('pter.item(6) == approx(1.15073)', locals())
```
failure_message: check(pter.item(6) == approx(1.15073))
success_message: check(pter.item(6) == approx(1.15073))


Tests for p1.5 (points: 5)
```
>>> check_str('by_pter.column(\'Date\').item(0) == \'2009-03-01\'', locals())
```
failure_message: check(by_pter.column('Date').item(0) == '2009-03-01')
success_message: check(by_pter.column('Date').item(0) == '2009-03-01')
```
>>> check_str('by_pter.column(\'PTER\').item(0) == approx(1.93946)', locals())
```
failure_message: check(by_pter.column('PTER').item(0) == approx(1.93946))
success_message: check(by_pter.column('PTER').item(0) == approx(1.93946))


Tests for p1.6 (points: 0)
```
>>> check_str('type(pter_over_time) == Table', locals())
```
failure_message: Be sure to store a new Table in the pter_over_time variable.
success_message: check(type(pter_over_time) == Table)
```
>>> check_str('pter_over_time.column("Date").item(0) == \'1994-03-01\'', locals())
```
failure_message: check(pter_over_time.column("Date").item(0) == '1994-03-01')
success_message: check(pter_over_time.column("Date").item(0) == '1994-03-01')


Tests for p1.7 (points: 5)
```
>>> check_str('type(highPTER) == bool', locals())
```
failure_message: check(type(highPTER) == bool)
success_message: check(type(highPTER) == bool)


Tests for p1.8 (points: 5)
```
>>> check_str('great_recession_quarters in between_or_equal(0, 16)', locals())
```
failure_message: check(great_recession_quarters in between_or_equal(0, 16))
success_message: check(great_recession_quarters in between_or_equal(0, 16))
```
>>> check_str('covid_quarters in between_or_equal(0, 8)', locals())
```
failure_message: check(covid_quarters in between_or_equal(0, 8))
success_message: check(covid_quarters in between_or_equal(0, 8))


Tests for p2.1 (points: 5)
```
>>> check_str('us_birth_rate in between_or_equal(0, 1)', locals())
```
failure_message: check(us_birth_rate in between_or_equal(0, 1))
success_message: check(us_birth_rate in between_or_equal(0, 1))


Tests for p2.2 (points: 5)
```
>>> check_str('movers.num_rows in between_or_equal(0, 52)', locals())
```
failure_message: check(movers.num_rows in between_or_equal(0, 52))
success_message: check(movers.num_rows in between_or_equal(0, 52))


Tests for p2.3 (points: 5)
```
>>> check_str('5e5 <= west_births <= 1e6', locals())
```
failure_message: check(5e5 <= west_births <= 1e6)
success_message: check(5e5 <= west_births <= 1e6)


Tests for p2.4 (points: 5)
```
>>> check_str('visualization in [1,2,3]', locals())
```
failure_message: visualization should be set to 1, 2, or 3.
success_message: check(visualization in [1,2,3])
```
>>> check_str('visualization == 3', locals())
```
failure_message: check(visualization == 3)
success_message: check(visualization == 3)


Tests for p2.6 (points: 5)
```
>>> check_str('type(assoc) == bool', locals())
```
failure_message: assoc should be set to True or False.
success_message: check(type(assoc) == bool)


Tests for p3.1 (points: 0)
```
>>> check_str('member_plot.get_title()==\'Member Rides\'', locals())
```
failure_message: check(member_plot.get_title()=='Member Rides')
success_message: check(member_plot.get_title()=='Member Rides')


Tests for p3.2 (points: 5)
```
>>> check_str('rides_between_20_and_25 in between_or_equal(5.5, 7.5)', locals())
```
failure_message: check(rides_between_20_and_25 in between_or_equal(5.5, 7.5))
success_message: check(rides_between_20_and_25 in between_or_equal(5.5, 7.5))


Tests for p3.3 (points: 0)
```
>>> check_str('casual_plot.get_title() == \'Casual Rides\'', locals())
```
failure_message: check(casual_plot.get_title() == 'Casual Rides')
success_message: check(casual_plot.get_title() == 'Casual Rides')
```
>>> check_str('casual_plot.get_ylim() == (0,0.07)', locals())
```
failure_message: check(casual_plot.get_ylim() == (0,0.07))
success_message: check(casual_plot.get_ylim() == (0,0.07))


Tests for p3.4 (points: 5)
```
>>> check_str('member_under_15 in between_or_equal(0, 100)', locals())
```
failure_message: check(member_under_15 in between_or_equal(0, 100))
success_message: check(member_under_15 in between_or_equal(0, 100))
```
>>> check_str('casual_under_15 in between_or_equal(0, 100)', locals())
```
failure_message: check(casual_under_15 in between_or_equal(0, 100))
success_message: check(casual_under_15 in between_or_equal(0, 100))


Tests for p3.5 (points: 5)
```
>>> check_str('member_median_bin in [1,2,3,4,5]', locals())
```
failure_message: check(member_median_bin in [1,2,3,4,5])
success_message: check(member_median_bin in [1,2,3,4,5])
