```python
# Initialize Otter
import otter
grader = otter.Notebook("pre02.ipynb")
```

<table style="width: 100%;">
<tr style="background-color: transparent;">
<td width="100px"><img src="https://cs104williams.github.io/assets/cs104-logo.png" width="90px" style="text-align: center"/></td>
<td>
  <p style="margin-bottom: 0px; text-align: left; font-size: 18pt;"><strong>CSCI 104: Data Science and Computing for All</strong><br>
                Williams College<br>
                Fall 2026</p>
</td>
</tr>


# Prelab 2: Tables

**Instructions**
- Before you begin, execute the cell at the TOP of the notebook to load the provided tests, as well as the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute these cells again.  
- Be sure to consult your [Python Reference](https://cs104williams.github.io/assets/python-library-ref.html)!
- Complete this notebook by filling in the cells provided. 
- Please be sure to not re-assign variables throughout the notebook.  For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously.
- There are no hidden tests in prelabs.

<hr/>
<h2>Setup</h2>


```python
# Run this cell to set up the notebook.
# These lines import the numpy, datascience, and cs104 libraries.

import numpy as np
from datascience import *
from cs104 import *
%matplotlib inline
```

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 1. Python Review (10 pts)



<font color='#B1008E'>
    
##### Learning objectives
- Practice using mathematical functions that are native to Python
</font>

The two building blocks of Python code are *expressions* and *statements*.  An **expression** is a piece of code that

* is self-contained, meaning it would make sense to write it on a line by itself, and
* usually evaluates to a value.

Here are two expressions that both evaluate to 3:

    3
    5 - 2
    
One important type of expression is the **call expression**. A call expression begins with the name of a function and is followed by the argument(s) of that function in parentheses. The function returns some value, based on its arguments. Some important mathematical functions are listed below.

| Function | Description                                                   |
|----------|---------------------------------------------------------------|
| `abs`      | Returns the absolute value of its argument                    |
| `max`      | Returns the maximum of all its arguments                      |
| `min`      | Returns the minimum of all its arguments                      |
| `pow`      | Raises its first argument to the power of its second argument |
| `round`    | Rounds its argument to the nearest integer                     |

Here are two call expressions that both evaluate to 3:

    abs(2 - 5)
    max(round(2.8), min(pow(2, 10), -1 * pow(2, 10)))

The expression `2 - 5` and the two call expressions given above are examples of **compound expressions**, meaning that they are actually combinations of several smaller expressions.  `2 - 5` combines the expressions `2` and `5` by subtraction.  In this case, `2` and `5` are called **subexpressions** because they're expressions that are part of a larger expression.

A **statement** is a whole line of code.  Some statements are just expressions.  The expressions listed above are examples.

Other statements *make something happen* rather than *having a value*. For example, an **assignment statement** assigns a value to a name. 

A good way to think about this is that we're **evaluating the right-hand side** of the equals sign and **assigning it to the left-hand side**. Here are some assignment statements:
    
    height = 1.3
    the_number_five = abs(-5)
    absolute_height_difference = abs(height - 1.688)

An important idea in programming is that large, interesting things can be built by combining many simple, uninteresting things.  The key to understanding a complicated piece of code is breaking it down into its simple components.

For example, a lot is going on in the last statement above, but it's really just a combination of a few things.  This picture describes what's going on.

<img src="statement.png" width="50%"/>

#### Part 1.1 (5 pts)


Hakeem is baking a cake. Recipe A calls for 225g grams of plain flour and Recipe B calls for 2.75 cups of flour. He wants to use the recipe that uses less flour. Hakeem converts both recipes to cups and then rounds to the nearest cup before he compares them. 

In the following cell, assign
- `recipe_a_cups` to the number of cups (rounded to the nearest whole cup) used by recipe A
- `recipe_b_cups` to the number of cups (rounded to the nearest whole cup) used by recipe B
- `min_recipe_cups` to the number of whole cups used by the recipe with the *fewest* cups 
- `recipe_with_least_flour` to a string, either "A" or "B," corresponding to the recipe with the fewest whole cups

Be sure to check your work by executing the test cell afterward.

*Hints:* 
- A cup of all-purpose flour typically weighs 120 grams
- Which functions (`abs`, `max`, `min`, `pow`, `round`) will you need to use? 

```python
recipe_a_cups = ...
recipe_b_cups = ...
min_recipe_cups = ...
recipe_with_least_flour = ...
```

```python
grader.check("p1.1")
```

#### Part 1.2 (5 pts)


In the next cell, assign the name `mystery` to the larger number among the following two numbers:

1. the **absolute value** of $2^{5}-2^{11}-2^1 + 1$ 
2. $5 \times 13 \times 31 + 11$.

Try to use just one statement (one line of code). Be sure to check your work by executing the test cell afterward.


```python
mystery = ...
mystery
```

```python
grader.check("p1.2")
```

We've asked you to use one line of code in the question above because it only involves mathematical operations and demonstrates a large, complex **statement**. 

However, more complicated programming questions will more require more steps. It isn’t always a good idea to jam these steps into a single line because it can make the code harder to read and harder to debug.

Good programming practice involves splitting up your code into smaller steps and using appropriate names. You'll have plenty of practice in the rest of this course!

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 2. Table Manipulation with National Park Data (30 pts)



<font color='#B1008E'>
    
##### Learning objectives
- Practice using built-in table functions: `.show(), .num_columns, .num_rows, .select(), .drop(), .sort(), .where()`
</font>

The table `species.csv` contains the [National Park Service's](https://irma.nps.gov/NPSpecies/) inventory of the animal species recorded in 56 U.S. national parks -- birds, mammals, insects, fish, reptiles, and more.  Each row represents one species observed in one park.

Run the next cell to load the `species` table.

```python
# Just run this cell

species = Table.read_table('species.csv')
```

Let's examine our table to see what data it contains.

#### Part 2.1: `.show()` (5 pts)


Use the method `show` to display the first 5 rows of `species`. 

*Note:* The terms "method" and "function" are technically not the same thing, but for the purposes of this course, we will use them interchangeably.

**Hint:** `tbl.show(3)` will show the first 3 rows of `tbl`. Additionally, make sure not to call `.show()` without an argument, as this may crash your kernel!

```python
...
```

```python
grader.check("p2.1")
```

Notice that some of the values in this table are missing, as denoted by `"nan."` This means either that the value is not available (e.g., if the park hasn't assessed how abundant a species is) or not applicable (e.g., `Conservation Status` is blank for species that are not endangered or threatened). You'll also notice that the table has a large number of columns in it

Before continuing, let's look up the `show` method in our online [Python Reference](http://cs104williams.github.io/assets/python-library-ref.html).  Go to that page and either scroll down the list of "Table Functions and Methods" to the entry for `show` or jump to it directly with the popup menu in the top right corner.  Familiarize yourself with the content of that row in the table, which describes how to use the method, the parameters it may take, and the result it computes.  Click on the row to see several simple examples of using the method.  Consult this documentation often!

#### Part 2.2: `.num_columns` and `.num_rows` (5 pts)


The table **property** `num_columns` returns the number of columns in a table. (A **property** is just a method that doesn't need to be called by adding parentheses.) 

Example call: `tbl.num_columns`

Look up `num_columns` in our online [Python Reference](http://cs104williams.github.io/assets/python-library-ref.html), and use `num_columns` to find the number of columns in our species dataset.  Assign the number of columns to `num_species_columns`.


```python
num_species_columns = ...
print("The table has", num_species_columns, "columns in it!")
```

```python
grader.check("p2.2")
```

Similarly, the property `num_rows` tells you how many rows are in a table.

```python
# Just run this cell

num_species_rows = species.num_rows
print("The table has", num_species_rows, "rows in it!")
```

#### Part 2.3: `.select()` (5 pts)


Several of the columns are bookkeeping about how each record was made -- whether it has been approved, whether the species is native, and so on.  If we're not interested in that information, it just makes the table difficult to read.  This comes up more than you might think, because people who collect and publish data may not know ahead of time what people will want to do with it.

In such situations, we can use the table method `select` to choose only the columns that we want in a particular table. It takes any number of arguments. Each should be the name of a column in the table. It returns a new table with only those columns in it. The columns are in the order *in which they were listed as arguments*.

For example, the value of `species.select("Park Name", "Category")` is a table with only the park and the category (bird, mammal, etc.) of each species in `species`.

Use `select` to create a table with only the park name, category, scientific name, and common names of each species.  Call that new table `species_names`.

*Hint:* Make sure to be exact when using column names with `select`; double-check capitalization and spaces!

```python
species_names = ...
species_names
```

```python
grader.check("p2.3")
```

#### Part 2.4:  `.drop()` (5 pts)


`drop` serves the same purpose as `select`, but it takes away the columns that you provide rather than the ones that you don't provide. Like `select`, `drop` returns a new table.

Suppose you just didn't want the `Species ID` and `Record Status` columns in `species`.  Create a table that's a copy of `species` but doesn't include those columns.  Call that table `species_without_id`.


```python
species_without_id = ...
species_without_id
```

```python
grader.check("p2.4")
```

#### Part 2.5: `.sort()` (5 pts)


Now, suppose we want to answer some questions about the parks themselves, rather than the species in them. The table `parks.csv` has one row per park, with its size and location.  Run the next cell to load it.

```python
# Just run this cell

parks = Table.read_table('parks.csv')
parks
```

For example, which park is farthest east?  Longitude increases as we move east, so let's sort `parks` by the `Longitude` column.

```python
parks.sort('Longitude')
```

Oops, that didn't answer our question because we sorted from smallest to greatest longitude, which put the *westernmost* park first. (Longitudes in the western hemisphere are negative.) To look at the greatest longitudes, we'll have to sort in reverse order.

```python
parks.sort('Longitude', descending=True)
```

The top row is now Acadia National Park in Maine -- the easternmost national park, and the closest one to Williamstown.

The `descending=True` bit is called a *named argument*. It has a default value of `False`, so when you explicitly tell the function `descending=True`, then the function will sort in descending order.


Some details about sort:

1. The first argument to `sort` is the name of a column to sort by.
2. If the column has text in it, `sort` will sort alphabetically; if the column has numbers, it will sort numerically.
3. The value of `parks.sort("Longitude")` is a *copy* of `parks`; the `parks` table doesn't get modified. For example, if we called `parks.sort("Longitude")`, then running `parks` by itself would still return the unsorted table.
4. Rows always stick together when a table is sorted.  It wouldn't make sense to sort just one column and leave the other columns alone.  For example, in this case, if we sorted just the `Longitude` column, the parks would all end up with the wrong longitudes.

Create a version of `parks` that's sorted by **latitude**, with the greatest latitudes first.  Call it `parks_by_latitude`.


```python
parks_by_latitude = ...
parks_by_latitude
```

```python
grader.check("p2.5")
```

#### Part 2.6: `.where()` (5 pts)


Now let's say we want a table of all the species recorded in Acadia National Park.  Sorting `species` by `Park Name` would put Acadia's rows at the top, but it wouldn't tell us where they end -- and it wouldn't help at all for a park in the middle of the alphabet.  Instead, we use the table method `where`.

```python
acadia_species = species.where('Park Name', are.equal_to('Acadia National Park'))
acadia_species
```

Let's dive into our use of `where` in more detail.  `where` takes 2 arguments:

1. The name of a column.  `where` finds rows where that column's values meet some criterion.
2. A **predicate** that describes the criterion that the column needs to meet. The predicate in the example above called the function `are.equal_to` with the value we wanted, 'Acadia National Park'.  We'll see other predicates soon.

The `where` method returns a new table that's a copy of the original table, but **with only the rows that meet the given predicate**.

The rows at the top of `acadia_species` are all mammals, but Acadia has plenty of other kinds of animals too.  The `sample` method picks rows at random from a table.  Run the next cell a few times to see species from different categories.

```python
acadia_species.sample(5)
```

Use `acadia_species` to create a table called `acadia_mammals` containing only the mammals recorded in Acadia.

```python
acadia_mammals = ...
acadia_mammals
```

```python
grader.check("p2.6")
```

So far we've only been using `where` with the predicate that requires finding the values in a column to be *exactly* equal to a certain value. However, there are many other predicates. See our [Python Reference](http://cs104williams.github.io/assets/python-library-ref.html#where) for a complete list.

<!-- BEGIN QUESTION -->

#### Part 2.7 (0 pts)


**This part is just for fun and practice -- it has no checks and is not graded.**

You now have two tables of real data and the whole toolbox from this section: `show`, `num_rows`, `select`, `drop`, `sort`, and `where`.  Pose a question of your own about the parks or the species in them, and try to answer it in the cell below.

Some ideas to get you started:
- Which park is the largest?  The smallest?
- How many bird species have been recorded in Yosemite National Park?
- Which parks are home to a species you find interesting -- moose, bald eagles, alligators?  (Try the predicate `are.containing` on the `Common Names` column.)
- What is the northernmost park outside of Alaska?

Write your question as a comment at the top of the cell, and then write the code that answers it.

```python
# Your question: ...
...
```

<!-- END QUESTION -->

<hr style="margin-bottom: 0px; padding:0; border: 2px solid #500082;"/>


## 3. Creating Tables (10 pts)



<font color='#B1008E'>
    
##### Learning objectives
- Create tables from arrays using `Table().with_columns()`
- Create tables from .csv files using `Table.read_table()`

So far, we've looked at and manipulated **tables**. But let's dig into what tables are and how we can create them. 

An **array** is useful for describing a single attribute of each element in a collection. For example, let's say our collection is all US States. Then an array could describe the land area of each state. 

Tables extend this idea by containing multiple columns stored and represented as arrays, each one describing a different attribute for every element of a collection. In this way, tables allow us to not only store data about many entities but to also contain several kinds of data about each entity. 

For example, in the cell below we have two arrays. The first one, `population_amounts`, contains the world population in each year from 2015 to 2023, as estimated by [Our World in Data](https://ourworldindata.org/population-growth). The second array, `years`, contains the years themselves. These elements are in order, so the year and the world population for that year have the same index in their corresponding arrays.

```python
# Just run this cell
population_amounts = make_array(7470491876, 7558554527, 7645617952, 7729902779, 7811293699,
                                7887001289, 7954448387, 8021407196, 8091734933)
years = np.arange(2015, 2023+1)
print("Population column:", population_amounts)
print("Years column:", years)
```

Suppose we want to answer this question:

> In which year did the world's population cross 8 billion?

You could technically answer this question just from staring at the arrays, but it's a bit convoluted, since you would have to count the position where the population first crossed 8 billion, then find the corresponding element in the years array. In cases like these, it might be easier to put the data into a `Table`. 

The expression below:

- creates an empty table using the expression `Table()`,
- adds two columns by calling `with_columns` with four arguments,
- assigns the result to the name `population`, and finally
- evaluates `population` so that we can see the table.

The strings `"Year"` and `"Population"` are column labels that we have chosen. The names `population_amounts` and `years` were assigned above to two arrays of the **same length**. The function `with_columns` (you can find the documentation [here](http://cs104williams.github.io/assets/python-library-ref.html)) takes in alternating strings (to represent column labels) and arrays (representing the data in those columns). The strings and arrays are separated by commas.

```python
population = Table().with_columns(
    "Year", years,
    "Population", population_amounts
)
population
```

Now the data is combined into a single table! It's much easier to parse this data. If you need to know what the population was in 2019, for example, you can tell from a single glance.

#### Part 3.1 (5 pts)


 In the cell below, we've created 2 arrays. Using the steps above, assign `top_10_movies` to a table that has two columns called "Name" and "Rating", which hold `top_10_movie_names` and `top_10_movie_ratings` respectively.



```python
top_10_movie_names = make_array(
        'The Shawshank Redemption (1994)',
        'The Godfather (1972)',
        'The Dark Knight (2008)',
        'The Lord of the Rings: The Return of the King (2003)',
        "Schindler's List (1993)",
        'The Godfather Part II (1974)',
        '12 Angry Men (1957)',
        'The Lord of the Rings: The Fellowship of the Ring (2001)',
        'Inception (2010)',
        'Fight Club (1999)')
top_10_movie_ratings = make_array(9.3, 9.2, 9.1, 9.0, 9.0, 9.0, 9.0, 8.9, 8.8, 8.8)

top_10_movies = ...

# We've put this next line here 
# so your table will get printed out 
# when you run this cell.
top_10_movies
```

```python
grader.check("p3.1")
```

#### Part 3.2 (5 pts)


In most cases, we aren't going to go through the trouble of typing in all the data manually. Instead, we load them in from an external source, like a data file. There are many formats for data files, but CSV ("comma-separated values") is the most common.

`Table.read_table(...)` takes one argument (a path to a data file in **string** format) and returns a table.  

The file `imdb.csv` contains a table of the 250 highest-rated movies on IMDb, among movies with at least 200,000 ratings, as of 2026.  Load it as a table called `imdb`.

```python
imdb = ...
imdb
```

```python
grader.check("p3.2")
```

Where did `imdb.csv` come from? Take a look at this lab's folder (in the left sidebar). You should see a file called `imdb.csv`.

Double-click to open up the `imdb.csv` file in that folder and look at the format. What do you notice? Jupyter displays the contents of the file as a basic table, but it is really just a specially-formatted text file.  In particular, the `.csv` filename ending says that this file is in the [CSV (comma-separated value) format](http://edoceo.com/utilitas/csv-file-format).

<hr class="m-0" style="border: 3px solid #500082;"/>

# You're Done!
Follow these steps to submit your work:
* Run the tests and verify that they pass as you expect. 
* Choose **Save Notebook** from the **File** menu.
* **Run the final cell** and click the link below to download the zip file. 

Once you have downloaded that file, go to [Gradescope](https://www.gradescope.com/) and submit the zip file to 
the corresponding assignment. For Prelab N, the assignment will be called "Prelab N Autograder".

Once you have submitted, your Gradescope assignment should show you passing all the tests you passed in your assignment notebook.


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for p1.1 (points: 5)
```
>>> check_str('recipe_a_cups == 2', locals())
```
failure_message: check(recipe_a_cups == 2)
success_message: check(recipe_a_cups == 2)
```
>>> check_str('recipe_b_cups == 3', locals())
```
failure_message: check(recipe_b_cups == 3)
success_message: check(recipe_b_cups == 3)
```
>>> check_str('min_recipe_cups == 2', locals())
```
failure_message: check(min_recipe_cups == 2)
success_message: check(min_recipe_cups == 2)
```
>>> check_str('recipe_with_least_flour == "A"', locals())
```
failure_message: check(recipe_with_least_flour == "A")
success_message: check(recipe_with_least_flour == "A")


Tests for p1.2 (points: 5)
```
>>> check_str('mystery == 2026', locals())
```
failure_message: check(mystery == 2026)
success_message: check(mystery == 2026)


Tests for p2.1 (points: 5)
```
>>> check_str('species != ...', locals())
```
failure_message: Make sure to run the cell above that assigns to `species`.
success_message: Looks good!


Tests for p2.2 (points: 5)
```
>>> check_str('num_species_columns == 13', locals())
```
failure_message: check(num_species_columns == 13)
success_message: check(num_species_columns == 13)


Tests for p2.3 (points: 5)
```
>>> check_str('np.sort(species_names.labels) == np.sort([\'Park Name\', \'Category\', \'Scientific Name\', \'Common Names\'])', locals())
```
failure_message: Be sure to select only the "Park Name", "Category", "Scientific Name", and "Common Names" columns!
success_message: Table has the correct labels!
```
>>> check_str('species_names.num_rows == 42570', locals())
```
failure_message: check(species_names.num_rows == 42570)
success_message: check(species_names.num_rows == 42570)


Tests for p2.4 (points: 5)
```
>>> check_str('species_without_id.num_columns == 11', locals())
```
failure_message: check(species_without_id.num_columns == 11)
success_message: check(species_without_id.num_columns == 11)
```
>>> check_str('np.sort(species_without_id.labels) == [ \'Abundance\', \'Category\', \'Common Names\', \'Conservation Status\', \'Family\', \'Nativeness\', \'Occurrence\', \'Order\', \'Park Name\', \'Scientific Name\', \'Seasonality\' ]', locals())
```
failure_message: Your table should have the following column labels: 'Abundance', 'Category', 'Common Names', 'Conservation Status', 'Family', 'Nativeness', 'Occurrence', 'Order', 'Park Name', 'Scientific Name', 'Seasonality'
success_message: Table has the correct labels!


Tests for p2.5 (points: 5)
```
>>> check_str('type(parks_by_latitude) == tables.Table', locals())
```
failure_message: check(type(parks_by_latitude) == tables.Table)
success_message: check(type(parks_by_latitude) == tables.Table)
```
>>> check_str('parks_by_latitude.column(\'Latitude\').take(range(3)) == [67.78, 67.55, 63.33]', locals())
```
failure_message: check(parks_by_latitude.column('Latitude').take(range(3)) == [67.78, 67.55, 63.33])
success_message: check(parks_by_latitude.column('Latitude').take(range(3)) == [67.78, 67.55, 63.33])


Tests for p2.6 (points: 5)
```
>>> check_str('acadia_mammals.num_rows == 55', locals())
```
failure_message: check(acadia_mammals.num_rows == 55)
success_message: check(acadia_mammals.num_rows == 55)
```
>>> check_str('all(acadia_mammals.column(\'Category\') == \'Mammal\')', locals())
```
failure_message: check(all(acadia_mammals.column('Category') == 'Mammal'))
success_message: check(all(acadia_mammals.column('Category') == 'Mammal'))


Tests for p3.1 (points: 5)
```
>>> check_str('top_10_movies.labels == make_array(\'Name\', \'Rating\')', locals())
```
failure_message: check(top_10_movies.labels == make_array('Name', 'Rating'))
success_message: check(top_10_movies.labels == make_array('Name', 'Rating'))
```
>>> check_str('top_10_movies.num_rows == 10', locals())
```
failure_message: check(top_10_movies.num_rows == 10)
success_message: check(top_10_movies.num_rows == 10)
```
>>> check_str('top_10_movies.column("Name").item(0) == \'The Shawshank Redemption (1994)\'', locals())
```
failure_message: check(top_10_movies.column("Name").item(0) == 'The Shawshank Redemption (1994)')
success_message: check(top_10_movies.column("Name").item(0) == 'The Shawshank Redemption (1994)')
```
>>> check_str('top_10_movies.column("Rating").item(0) == 9.3', locals())
```
failure_message: check(top_10_movies.column("Rating").item(0) == 9.3)
success_message: check(top_10_movies.column("Rating").item(0) == 9.3)


Tests for p3.2 (points: 5)
```
>>> check_str('imdb.labels == make_array(\'Title\', \'Year\', \'Rating\', \'Votes\')', locals())
```
failure_message: check(imdb.labels == make_array('Title', 'Year', 'Rating', 'Votes'))
success_message: check(imdb.labels == make_array('Title', 'Year', 'Rating', 'Votes'))
```
>>> check_str('imdb.row(0).item("Votes") == 3208129', locals())
```
failure_message: check(imdb.row(0).item("Votes") == 3208129)
success_message: check(imdb.row(0).item("Votes") == 3208129)
```
>>> check_str('imdb.row(0).item("Year") == 1994', locals())
```
failure_message: check(imdb.row(0).item("Year") == 1994)
success_message: check(imdb.row(0).item("Year") == 1994)
